evaluation_interval: 1m

rule_files:
  - build_service_availability_recording_rules.yaml

tests:
  - interval: 1m
    name: RSExporterTest
    input_series:
      - series: "redhat_appstudio_buildservice_global_github_app_available{service='build-service-controller-manager-metrics-service', check='github'}"
        values: "1 1 1 1 1"
      - series: "redhat_appstudio_buildservice_global_github_app_available{service='build-service-controller-manager-metrics-service1', check='github'}"
        values: "0 0 0 0 0"
      - series: "redhat_appstudio_buildservice_global_github_app_available{service='build-service-controller-manager-metrics-service2', check='github'}"
        values: "0 1 0 1 0"
    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='build-service-controller-manager-metrics-service', check='github'}
            value: 1
          - labels: konflux_up{service='build-service-controller-manager-metrics-service1', check='github'}
            value: 0
          - labels: konflux_up{service='build-service-controller-manager-metrics-service2', check='github'}
            value: 0
