evaluation_interval: 1m

rule_files:
- exporter_recording_rules.yml

tests:
- interval: 1m
  name: DSExporterTest
  input_series:
    - series: "grafana_ds_up{check='appstudio-ds', label='label'}"
      values: "1 1 1 1 1"
    - series: "grafana_ds_up{check='github', label='label'}"
      values: "0 0 0 0 0"
    - series: "grafana_ds_up{other_label='sso', label='label'}"
      values: "0 1 0 1 0"
  promql_expr_test:
    - expr: konflux_up
      eval_time: 5m
      exp_samples:
        - labels: konflux_up{service='grafana', check='appstudio-ds', label='label'}
          value: 1
        - labels: konflux_up{service='grafana', check='github', label='label'}
          value: 0
        - labels: konflux_up{service='grafana', other_label='sso', label='label'}
          value: 0
